# Assignment 1

For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.

Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

```python
example_word_count()
```

Output:
```
5
```

## Part A

Find a list of all of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    return list(re.findall("[A-Z][a-z]*\w",simple_string))

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re

def grades():
    with open("assets/grades.txt", "r") as file:
        grades = file.read()  # Read the contents of the file.
        
    pattern = r"^.*: B"  # Define a regular expression pattern to match lines ending in ": B".
    
    students = re.findall(pattern, grades, re.MULTILINE)  # Find all matching lines in the file.
    
    split_lst = [elem.split(':')[0] for elem in students]  # Extract the student names by splitting on the ":" character.
    
    return list(split_lst)  # Return the list of student names.

```

```python
grades()
```

Output:
```
['Bell Kassulke',
 'Simon Loidl',
 'Elias Jovanovic',
 'Hakim Botros',
 'Emilie Lorentsen',
 'Jake Wood',
 'Fatemeh Akhtar',
 'Kim Weston',
 'Yasmin Dar',
 'Viswamitra Upandhye',
 'Killian Kaufman',
 'Elwood Page',
 'Elodie Booker',
 'Adnan Chen',
 'Hank Spinka',
 'Hannah Bayer']
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:

* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re

def logs():
    data = []  # Initialize an empty list to store the extracted log data.
    with open("assets/logdata.txt", "r") as f:  # Open the log data file.
        logdata = f.read()  # Read the contents of the log data file.
        for item in re.finditer(
            r'(?P<host>(?:\d+\.){1,3}\d+)\s+-\s+'  # Match the IP address.
            r'(?P<user_name>[\w+\-]+)?\s+'  # Match the user name.
            r'\[(?P<time>[-\w\s:/]+)\]\s+'  # Match the timestamp.
            r'"(?P<request>\w+.+?)"',  # Match the HTTP request.
            logdata  # Search for matches in the log data.
        ):
            x = item.groupdict()  # Extract the named groups from the match as a dictionary.
            if x["user_name"] is None:  # If the user name is missing, replace it with a dash.
                x["user_name"] = "-"
            data.append(x)  # Add the extracted log data to the list.
    return data  # Return the list of extracted log data.
```


(?P<host>(?:\d+\.){1,3}\d+): This pattern matches an IP address in the form xxx.xxx.xxx.xxx, where each xxx is a number between 0 and 255. The pattern uses a non-capturing group (?:\d+\.) to match one or more digits followed by a period, and then repeats that group between 1 and 3 times with the syntax {1,3}. The final \d+ matches the last octet of the IP address. The (?P<host>...) syntax creates a named capturing group with the name host.

\s+-\s+: This pattern matches the separator between the IP address and the user name, which is a dash surrounded by one or more whitespace characters.

(?P<user_name>[\w+\-]+)?\s+: This pattern matches the user name, which may be missing or may contain one or more alphanumeric characters, underscores, or dashes. The ? makes the pattern optional. The (?P<user_name>...) syntax creates a named capturing group with the name user_name.

\[(?P<time>[-\w\s:/]+)\]\s+: This pattern matches the timestamp enclosed in square brackets, which includes a date and time in the format dd/mm/yyyy:hh:mm:ss -xxxx. The pattern matches any combination of alphanumeric characters, spaces, dashes, colons, and slashes inside the square brackets. The (?P<time>...) syntax creates a named capturing group with the name time.

"(?P<request>\w+.+?)": This pattern matches the HTTP request enclosed in quotes, which typically consists of a method (such as GET or POST) and a URL. The pattern matches any alphanumeric characters in the method, followed by one or more characters of any type (matched by .+?) in the URL. The (?P<request>...) syntax creates a named capturing group with the name request.

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

```python

```

```python

```

```python

```